---
title: '(OLD)Test your data pipeline'
description: 'Before, after or while you build your data pipeline you want to test it and gain confidence in its behaviour. '
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Writing unit tests for pyspark

```yaml
type: VideoExercise
key: 7f879913b2
xp: 50
```

`@projector_key`
de549e4812c9bd9e7bb4c9fce8e9dc4f
